# Assignment 1 (Total Points: 10) <a class="anchor" id="top"></a>


## Submission instructions:

- The assignment is designed for teams of two students.
- Add your solutions to the cells with the following comment `# YOUR CODE HERE`, which denotes that these sections are graded.
- The notebook that you will submit has to have the student ids, separated by underscores (E.g., `12341234_12341234.ipynb`). 
- This will be parsed by a regexp, so please double check your filename.
- Only one member of each group should submit the notebook file to canvas.
- Make sure to check that your notebook runs before submission. A quick way to do this is to restart the kernel and run all the cells.  
- Please do not delete cells. Removing cells can lead to grade deduction. You are allowed to add cells for performing tests or adding auxiliary methods in the given cells.
- Note, that you are not allowed to use Google Colab.
- Add your name and e-mail in the following cell:

```python
## Add your information in here.
NAME = ""
NAME2 = ""
EMAIL = ""
EMAIL2 = ""
```

**Assignment instructions Goals**:

In this assignment, you will program several algorithms related to BombKnightSudoku (a variant of Sudoku) on a $(k*k)^2$ size grid, for varying sizes of $k$. You will use several problem solving and search methods.

The homework assignment consists of several parts (a)–(d). You get 1 point for free. For the rest parts you need to develop the following algorithms for solving sudoku game:

 - Part (a) 3pts: solving puzzles with propagation
 - Part (b) 1pt as a bonus: solving puzzles using SAT
 - Part (c) 4pts: solving puzzles using ASP
 - Part (d) 2pts: generating puzzles with unique solutions using ASP (or you may opt to solve an easier variant for a maximum of 1 point)

    
**Resources**: 

- Literature:
- Russell, S., & Norvig, P. (2002). Artificial intelligence: a modern approach. [http://aima.cs.berkeley.edu/]
- Gebser, M., Kaminski, R., Kaufmann, B., & Schaub, T. (2012). Answer set solving in practice. Synthesis lectures on artificial intelligence and machine learning, 6(3), 1-238. [https://potassco.org/book/]
- Van Harmelen, F., Lifschitz, V., & Porter, B. (Eds.). (2008). Handbook of knowledge representation. Elsevier.

**Game information: BombKnightSudoku's**

Here follows a description of inputs and solutions for **BombKnightSudoku's**. These notions are similar (but not exactly the same) as for the Sudoku's that you might have seen elsewhere, generalized to arbitrary values of k (typically, $k=3$).

An input consists of a $k$-by-$k$ grid of blocks, where each consisting of a $k$-by-$k$ grid of cells, forming a $k*k$-by-$k*k$ grid of cells. Each cell can (but must not) contain a value between 1 and $k*k$.

An example is the following grid, for $k$=3, written out as $k*k$ lines, each consisting of $k*k$ numbers, separated by spaces. Here the value $0$ represents that a cell does not contain a value between $1$ and $k*k$.

```
0 0 0 0 0 0 0 0 0
0 3 0 0 8 0 0 0 0
0 0 0 0 0 0 0 0 0
0 0 0 0 4 0 0 0 5
0 0 0 0 0 0 0 2 7
0 0 0 0 0 0 3 0 0
0 0 0 0 0 0 9 0 0
0 0 0 0 5 6 0 0 0
0 0 0 0 0 0 0 0 0
```

Here is the same input, where now the $k$-by-$k$ blocks are separated from each other with whitespace.

```
0 0 0  0 0 0  0 0 0
0 3 0  0 8 0  0 0 0
0 0 0  0 0 0  0 0 0

0 0 0  0 4 0  0 0 5
0 0 0  0 0 0  0 2 7
0 0 0  0 0 0  3 0 0

0 0 0  0 0 0  9 0 0
0 0 0  0 5 6  0 0 0
0 0 0  0 0 0  0 0 0
```

A solution for a sudoku input is a $k*k$-by-$k*k$ grid (for the same $k$), where:

 - Each cell contains a value between 1 and $k*k$.
If a cell ($i$,$j$) contains a value u in the input, then the cell ($i$,$j$) in the solution must contain the same value u.
 - Each two different cells in the same row must contain different values.
 - Each two different cells in the same column must contain different values.
 - Each two different cells in the same $k*k$ block must contain different values.
 - Each two different cells that are directly adjacent (horizontally, diagonally or vertically) must contain different values. In other words, a cell cannot contain the same value as any of the eight cells surrounding it.
 - Each two different cells that can be reached from each other with a single knight's move (as in Chess) must contain different values.
(The last two constraints are what distinguish BombKnightSudoku's from regular Sudoku's.)

For example, a solution for the input mentioned above is the following:

```
5 6 2 7 3 4 1 8 9
7 3 4 1 8 9 5 6 2
1 8 9 5 6 2 7 3 4
6 2 7 3 4 1 8 9 5
3 4 1 8 9 5 6 2 7
8 9 5 6 2 7 3 4 1
2 7 3 4 1 8 9 5 6
4 1 8 9 5 6 2 7 3
9 5 6 2 7 3 4 1 8
```

**Representation of inputs and outputs**

Inputs are represented as a list of lists of numbers.
For example, the BombKnightSudoku input above is represented as:

```
input =
    [[0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 3, 0, 0, 8, 0, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 0, 0, 0, 4, 0, 0, 0, 5],
     [0, 0, 0, 0, 0, 0, 0, 2, 7],
     [0, 0, 0, 0, 0, 0, 3, 0, 0],
     [0, 0, 0, 0, 0, 0, 9, 0, 0],
     [0, 0, 0, 0, 5, 6, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0]]
```

Solutions are represented similarly, for example:

```
solution =
    [[5, 6, 2, 7, 3, 4, 1, 8, 9],
     [7, 3, 4, 1, 8, 9, 5, 6, 2],
     [1, 8, 9, 5, 6, 2, 7, 3, 4],
     [6, 2, 7, 3, 4, 1, 8, 9, 5],
     [3, 4, 1, 8, 9, 5, 6, 2, 7],
     [8, 9, 5, 6, 2, 7, 3, 4, 1],
     [2, 7, 3, 4, 1, 8, 9, 5, 6],
     [4, 1, 8, 9, 5, 6, 2, 7, 3],
     [9, 5, 6, 2, 7, 3, 4, 1, 8]]
```

Note that some test_inputs are given in the auxiliary python file which are used for testing your methods.

The necessary libraries for developing all necessary code can be found in the following block:

```python
# Necessary libraries that will be employed
import itertools

from copy import deepcopy
from dataclasses import dataclass
from typing import List
from sudoku_func import *
#!{sys.executable} -m pip install clingo
#!{sys.executable} -m pip install python-sat
#!{sys.executable} -m pip install codetiming
import clingo
import pysat
from pysat.formula import CNF
from pysat.solvers import MinisatGH
```

## (a): solving puzzles with propagation

Implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use an *exhaustive search with propagation (inference)* approach (see Section 6.3 from the AI book by Russell and Norvig). That is, it should traverse a search tree, and it should perform propagation at each of the nodes in the search tree.

Before starting coding the core part of this algorithm, you should develop several methods that will help you coding the propagation (for instance a method for finding uncertain cells or a function that checks if you have found a solution, a method that checks for possible contradictions and so forth). 

Finally, the code that calls propagation recursively is given later.

Firstly, you should define a function that given a list with possible values at the current state of the search returns the first found cell where there is still more than one possibility.

* Inputs: 
    * sudoku_possible_values: is a list containing all the possible values at the current state of the search. This variable is a 3d list (list(list(list(int)))) since contains the possible values for all input puzzles (test_inputs from sudoku_func.py).
    * k: the size of the grid.
* Output:
    * the first-found cell.  

<!-- grade_id: cell-821a87125f0afd92, points: 0, grade: false, solution: true -->
```python
def find_uncertain_cell(sudoku_possible_values, k):    
        """
        Finds a cell where there is still more than one possibility.
        Inputs: 
            - sudoku_possible_values: all the possible values from the current step of the game.
            - k: the size of the grid.
        Output:
            - the found cell.     
        """

        # YOUR CODE HERE
        raise NotImplementedError()
```

Next, you should define a function that checks whether you have found a solution or not which should be called at each current state of the search.

* Input: 
    * sudoku_possible_values: is a list containing all the possible values at the current state of the search. This variable is a 3d list (list(list(list(int)))) since contains the possible values for all input puzzles (test_inputs from sudoku_func.py).
* Output:
    * A boolean variable which takes True if a solution is found or False otherwise.

<!-- grade_id: cell-1dee4c4c3585d194, points: 0, grade: false, solution: true -->
```python
def found_solution(sudoku_possible_values):
        """
        Checks if the search found a solution.
        Input: 
            - sudoku_possible_values: all the possible values from the current step of the game.
        Output:
            - A boolean variable which takes True if a solution is found or False otherwise.
        """
        # YOUR CODE HERE
        raise NotImplementedError()
```

Then, you should fill the following function with all possible contradictions that you should take into account.

* Inputs: 
    * sudoku_possible_values: is a list containing all the possible values at the current state of the search. This variable is a 3d list (list(list(list(int)))) since contains the possible values for all input puzzles (test_inputs from sudoku_func.py).
    * k: the size of the grid.
* Output:
    * A boolean variable: True if we run into a contradiction and False otherwise.
    
For instance, the first contradiction is the following: Contradiction type 1: some cell has no further possible values. If yes then you should return True. 

Then, you should come up with extra contradictions.

<!-- grade_id: cell-158784836524a39d, points: 0, grade: false, solution: true -->
```python
def contradiction(sudoku_possible_values, k):            
        
        """
        Checks if the search ran into a contradiction.
        Inputs: 
            - sudoku_possible_values: all the possible values from the current step of the game.
            - k: the size of the grid.
        Output:
            - A boolean variable: True if we run into a contradiction and False otherwise.
        """
        
        # Contradiction type 1: some cell has no further possible values. If yes then you should return True.
        # Further contradictions. (What other contradictions should you consider here?)”
        
        # YOUR CODE HERE
        raise NotImplementedError()
        # Otherwise: no contradiction is found
        return False
```

Next, you should fill propagate function where you will add the functionality of the solver. The goal is to remove the possibilities in a tile that contradict the sudoku rules.

* Inputs: 
    * sudoku_possible_values: is a list containing all the possible values at the current state of the search. This variable is a 3d list (list(list(list(int)))) since contains the possible values for all input puzzles (test_inputs from sudoku_func.py).
    * k: the size of the grid.
* Output:
    * sudoku_possible_values: The updated puzzle after the application of propagation.

<!-- grade_id: cell-3415cc791e17a8cc, points: 0, grade: false, solution: true -->
```python
def propagate(sudoku_possible_values, k):
        """
        Carry out propagation.
        Inputs: 
            - sudoku_possible_values: all the possible values from the current step of the game.
            - k: the size of the grid.
        Output:
            - sudoku_possible_values: The updated puzzle after the application of propagation.
        """
    
        # YOUR CODE HERE
        raise NotImplementedError()
```

The following method is a recursive function to solve the sudoku, using propagation.

* Inputs: 
    * sudoku_possible_values: is a list containing all the possible values at the current state of the search. This variable is a 3d list (list(list(list(int)))) since contains the possible values for all input puzzles (test_inputs from sudoku_func.py).
    * k the size of the grid.
* Output:
    *  Returns a generator with the solved puzzle. More info about the generators and the way that they work can be found in the following [link](https://stackabuse.com/python-generators).

```python
def solve_recursively(sudoku_possible_values, k):
    
        """
        Recursive function to solve the sudoku, using propagation
        Inputs: 
            - sudoku_possible_values: all the possible values from the current step of the game.
            - k: the size of the grid.
        Output:
            - Returns a generator with the solved puzzle.
        """
        # Check if we ran into a contradiction.
        if contradiction(sudoku_possible_values, k):
            return []
        # Check if we have a valid solution
        if found_solution(sudoku_possible_values):
            yield sudoku_possible_values
        else:
            # Propagate
            sudoku_possible_values = propagate(sudoku_possible_values, k)
            # Check for contradictions
            if contradiction(sudoku_possible_values, k):
                return []
            if found_solution(sudoku_possible_values):
                yield sudoku_possible_values
            else:
                # Find a cell that is still uncertain
                uncertain_cell = find_uncertain_cell(sudoku_possible_values, k)
                (row, col) = uncertain_cell
                # Recurse on the different values for cell (row, col)
                possibilities = sudoku_possible_values[row][col]
                for possibility in possibilities:
                    sudoku_possible_values_copy = deepcopy(
                        sudoku_possible_values
                    )
                    sudoku_possible_values_copy[row][col] = [possibility]
                    for solution in solve_recursively(
                            sudoku_possible_values_copy,
                            k
                    ):
                        yield solution                
```

Finally, call the recursive method from the previous cell to solve the game using propagation.

Carry out constraint propagation (as part of a backtracking search algorithm for solving a sudoku input), by removing possible values for some cells based on the remaining possible values for other cells.

* Parameters:
    * sudoku (list(list(list(int)))): A list containing all the input puzzles (found in sudoku_func.py)
    * k (int): The dimension of the sudoku input.

* Returns:
    * A generator containing the solved puzzle.

```python
def solve_sudoku_propagate(sudoku,k):
    """
    Carry out constraint propagation (as part of a backtracking search algorithm for solving a sudoku input), by removing possible values for some cells based on the remaining possible values for other cells.
    Inputs:
        sudoku (list(list(list(int)))): A list containing all the input puzzles (found in sudoku_func.py)
        k (int): The dimension of the sudoku input.

    Outpus:
        A generator that contain the solved puzzle.
    """
    # Initialize the data structure of sudoku_possible_values
    sudoku_possible_values = []
    for row in sudoku:
        row_possibilities = []
        for element in row:
            if element == 0:
                possibilities = list(range(1, k**2+1))
            else:
                possibilities = [element]
            row_possibilities.append(possibilities)
        sudoku_possible_values.append(row_possibilities)

    sudoku_possible_values = propagate(sudoku_possible_values, k)

    # Solve the sudoku by recursion
    for solution in solve_recursively(sudoku_possible_values, k):
        # Transform the data structure into a solution, and yield it
        solved_sudoku = [
            [solution[row][col][0] for col in range(0, k**2)]
            for row in range(0, k**2)
        ]
        yield solved_sudoku
```

The following code meants to test whether your provided solution has the same number of possibilities as the groundtruth for the puzzles.

<!-- grade_id: cell-61d968bd0a091d73, points: 3, grade: true, solution: false -->
```python
# Automatic checking of your code. Calls the generator and checks the solutions.
from sudoku_func import check_num_solutions
from sudoku_func import test_inputs
solver = solve_sudoku_propagate
## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(test_input.sudoku, test_input.k, test_input.num_solutions,solver) == True
    print("  Found exactly the set of correct solutions")
```

## (b): solving puzzles using SAT (bonus part)

Again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into a propositional CNF formula &varphi;, calling a SAT solver to find satisfying assignments for &varphi;, and translating these assignments to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-940a32cfcc687489, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_SAT(sudoku,k):
    """
    Solve a sudoku input by encoding the problem into SAT, calling a SAT solver, and retrieving the solution for the sudoku input from a satisfying truth assignment given by the SAT solver.

    Parameters:
        sudoku (list(list(list(int)))): A list containing all the input puzzles (found in sudoku_func.py)
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

The following code meants to test whether your provided solution has the same number of possibilities as the groundtruth for the puzzles.

<!-- grade_id: cell-4f0c31f4bfc245aa, points: 1, grade: true, solution: false -->
```python
# Automatic checking of your code. Calls the generator and checks the solutions. 
from sudoku_func import check_num_solutions
from sudoku_func import test_inputs
solver = solve_sudoku_SAT

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(test_input.sudoku, test_input.k, test_input.num_solutions,solver) == True
    print("  Found exactly the set of correct solutions")
```

## (c): solving puzzles using ASP

Yet again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-e08fbc765eba6e6d, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_ASP(sudoku,k):
    """
    Solve a sudoku input by encoding the problem into ASP, calling an ASP solver, and retrieving the solution for the sudoku input from an answer set given by the ASP solver.

    Parameters:
        sudoku (list(list(list(int)))): A list containing all the input puzzles (found in sudoku_func.py)
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

The following code meants to test whether your provided solution has the same number of possibilities as the groundtruth for the puzzles.

<!-- grade_id: cell-d58b222e808c0f0e, points: 4, grade: true, solution: false -->
```python
# Automatic checking of your code. Calls the generator and checks the solutions.
from sudoku_func import check_num_solutions
from sudoku_func import test_inputs
solver = solve_sudoku_ASP

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(test_input.sudoku, test_input.k, test_input.num_solutions,solver) == True
    print("  Found exactly the set of correct solutions")
```

## (d): generating puzzles with unique solutions using ASP

Implement an algorithm that generates `9`-by-`9` BombKnightSudoku inputs (i.e., `k=3`) that:
1. have exactly 10 cells filled with a number between `1` and `9`,
1. and have a unique solution.
That is, your algorithm should output a generator that iterates over BombKnightSudoku inputs with the required properties.

Your algorithm should use the approach of translating this problem to an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to BombKnightSudoku inputs with the required properties.

(To be clear: you may not 'hardcode' a list of BombKnightSudoku inputs in the answer set program *P*.)

There are two variants of this assignment: a harder variant (d.I) and an easier variant (d.II). You may choose which variant you would like to solve.
If you opt for the harder variant (d.I), you can get a maximum of 3 points for this assignment.
If you opt for the easier variant (d.II), you can get a maximum of 2 points for this assignment.


### (d.I) Harder variant (2pts):

To solve the harder variant, your answer set program *P* should be such that for **each** BombKnightSudoku input *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you should express in *P* the constraints 1. (that exactly 10 cells are filled) and 2. (that the input has a unique solution).

*Hint:* Use the technique of saturation to express constraint 2.

### (d.II) Easier variant (1pt):

To solve the easier variant, your answer set program *P* only needs to be such that for **a subset of** BombKnightSudoku inputs *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you could express in *P*, in addition to the constraint 1. (that exactly 10 cells are filled), some other property that entails constraint 2. (that the input has a unique solution).

*Hint:* If a Sudoku input can be solved using some [basic strategies](http://pi.math.cornell.edu/~mec/Summer2009/Mahmood/Solve.html), it must have a unique solution. Consider variants of these basic strategies for BombKnightSudoku.

<!-- grade_id: cell-fb57584e65e58ba1, points: 0, grade: false, solution: true -->
```python
def check_generator():
    """
    Checks whether the BombKnightSudoku puzzle generator returns puzzles that
    have a unique solution *and* have exactly 10 cells filled in.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d86460df48d77382, points: 2, grade: true, solution: false -->
```python
# Test whether your function returns sudoku puzzles with unique solutions.

```